# IS319 - Deep Learning

## TP1 - Neural networks

The goal of this TP is to implement a simple feedforward neural network, but without the use of libraries like PyTorch or TensorFlow. We will only use NumPy.

```python
import numpy as np
```

### 1. Activation function and its derivative

**(Question)** Implement the following activation function and its respective gradient (vector of partial derivatives). These should be applied element-wise to the input vector `a`.

<!-- grade_id: cell-24345ebc30158580, points: 0, grade: false, solution: true -->
```python
def sigmoid(a):
    '''Return the element-wise sigmoid of the input vector.'''
    return 1/(1+np.exp(-a))

def d_sigmoid(a):
    '''Return the partial derivatives of the sigmoid function
    with respect to the input vector.'''
    return np.multiply(sigmoid(a),(np.ones_like(sigmoid(a)) - sigmoid(a)))
```

<!-- grade_id: activation-functions, points: 0, grade: true, solution: false -->
```python
a = np.random.randn(100)
assert np.all(sigmoid(a) >= 0.)
assert np.all(sigmoid(a) <= 1.)
assert sigmoid(0.) == 0.5
assert np.all(d_sigmoid(a) >= 0.)
assert np.all(d_sigmoid(a) <= 0.25)
```

### 2. Loss function and its derivative

**(Question)** Implement the following loss function and its respective gradient (vector of partial derivatives).

`y` and `d` correspond to predictions and ground-truth labels respectively. They are assumed to be be matrices of size `n_classes * n_samples`.

<!-- grade_id: cell-0570574f9ec3aec5, points: 0, grade: false, solution: true -->
```python
def squared_error(y, d):
    '''Return a scalar corresponding to the sum of squared errors.'''
    # The sum instead of mean will be more convenient for this TP
    return 0.5*np.sum(np.power(y-d, 2))
def d_squared_error(y, d):
    '''Return the vector of partial derivatives of the sum of
    squared errors with respect to the predictions.'''
    return y-d
```

<!-- grade_id: loss-function, points: 0, grade: true, solution: false -->
```python
y = np.random.randn(3, 100)
d = np.random.randn(3, 100)
assert squared_error(y, d) >= 0.
assert d_squared_error(y, d).shape == y.shape
```

### 3. Neural network architecture

We will implement a simple fully-connected neural network with **one hidden layer** and **one output layer**.

This neural network is defined by a number of inputs, a number of hidden units, and a number of output units.

The activation function will be sigmoid and the loss function will be the sum of squared errors, both implemented above.

**(Question)** Complete the class below to initialize the weights and biases randomly.

<!-- grade_id: cell-4ae365acb7330ae3, points: 0, grade: false, solution: true -->
```python
class NeuralNetwork():
    def __init__(self, n_input, n_hidden, n_output):
        '''Initialize a neural network with `n_input` input neurons,
        `n_hidden` hidden neurons and `n_output` output neurons.'''
        self.n_input = n_input
        self.n_hidden = n_hidden
        self.n_output = n_output
        self.init_weights()
        
    def init_weights(self):
        '''Initialize random weights with attributes `W1`, `b1`, `W2` and `b2`.'''
        self.W1 = np.random.rand(self.n_hidden, self.n_input)
        self.b1 = np.random.rand(self.n_hidden, 1)
        self.W2 = np.random.rand(self.n_output, self.n_hidden)
        self.b2 = np.random.rand(self.n_output, 1)
```

<!-- grade_id: init-weights, points: 0, grade: true, solution: false -->
```python
nn = NeuralNetwork(64, 32, 3)
assert nn.W1.ndim == 2
assert nn.b1.ndim == 2
assert nn.W2.ndim == 2
assert nn.b2.ndim == 2
```

### 4. Forward pass

The forward pass is defined as:
$$\begin{align*}
\mathbf{h}_1 &= \sigma(\mathbf{a}_1) \quad\text{with}\quad \mathbf{a}_1 = \mathbf{W}_1 \mathbf{x} + \mathbf{b}_1 \\
\mathbf{y} &= \sigma(\mathbf{a}_2) \quad\text{with}\quad \mathbf{a}_2 = \mathbf{W}_2 \mathbf{h}_1 + \mathbf{b}_2
\end{align*}$$

**(Question)** Implement the forward pass for input examples `X`. Save intermediate results `a1`, `h1` and `a2` into attributes (as they will be needed for the backpropagation algorithm).

<!-- grade_id: cell-ec6cc8adc2e96480, points: 0, grade: false, solution: true -->
```python
class NeuralNetwork(NeuralNetwork): # (the method will be added to the `NeuralNetwork` class)
    def forward(self, X):
        self.a1 = self.W1@X + self.b1
        self.h1 = sigmoid(self.a1)
        self.a2 = self.W2@self.h1 + self.b2
        return sigmoid(self.a2)
```

<!-- grade_id: forward-pass, points: 0, grade: true, solution: false -->
```python
nn = NeuralNetwork(64, 32, 3)
X = np.random.randn(64, 100)
y = nn.forward(X)
assert y.shape == (3, 100)
```

**(Question)** Implement the function below to obtain a classification decision from the network. To do that, apply the forward pass, then choose the class corresponding to the maximum output value.

*Hint:* use the `np.argmax` function.

<!-- grade_id: cell-07e51093c0c43eef, points: 0, grade: false, solution: true -->
```python
class NeuralNetwork(NeuralNetwork): # (the method will be added to the `NeuralNetwork` class)
    def predict(self, X):
        return np.argmax(self.forward(X), axis = 0)
```

<!-- grade_id: predict, points: 0, grade: true, solution: false -->
```python
nn = NeuralNetwork(64, 32, 3)
X = np.random.randn(64, 100)
y = nn.predict(X)
assert y.shape == (100,)
assert np.any(y == 0) or np.any(y == 1) or np.any(y == 2)
```

### 5. Backward pass

**(Question)** Implement the backward pass for input examples `X`, ground-truth `d`, predictions `y`.

*Advice:* keep track of the shapes of each derivative using comments.

<!-- grade_id: cell-dc848db6c6a52963, points: 0, grade: false, solution: true -->
```python
class NeuralNetwork(NeuralNetwork):
    def backward(self, X, y, d):
        '''Compute the partial derivatives of the loss function
        with respect to all weights of the neural network.
        Return these in variables `d_W1`, `d_b1`, `d_W2` and `d_b2`.'''
        # Backpropagation for the output layer
        # You should compute, d_ey, d_ya2, d_a2w2 and finally delta2
        # Then you can compute d_W2 and d_b2
        d_ey = d_squared_error(y, d)
        d_ya2 = d_sigmoid(self.a2)
        d_a2w2 = self.h1
        delta2 = np.multiply(d_ey, d_ya2)
        d_W2 = delta2@d_a2w2.T
        d_b2 = np.sum(delta2, axis=1, keepdims=True)

        
        # Backpropagation for the hidden layer
        # You should comput d_h1a1 and finally delta1
        # Then you can compute d_W1 and d_b1
        d_h1a1 = d_sigmoid(self.a1)
        d_a1w1 = X
        delta1 = np.multiply(self.W2.T@delta2, d_h1a1)
        d_W1 = delta1@d_a1w1.T
        d_b1 = np.sum(delta1, axis=1, keepdims=True)

        
        return d_W1, d_b1, d_W2, d_b2
```

<!-- grade_id: backward, points: 0, grade: true, solution: false -->
```python
nn = NeuralNetwork(64, 32, 3)
X = np.random.randn(64, 100)
d = np.random.randint(0, 2, size=(3, 100))
y = nn.forward(X)
loss = squared_error(y, d)
d_W1, d_b1, d_W2, d_b2 = nn.backward(X, y, d)
assert d_W1.shape == nn.W1.shape
assert d_b1.shape == nn.b1.shape
assert d_W2.shape == nn.W2.shape
assert d_b2.shape == nn.b2.shape
```

### 6. Weights update with gradient descent

**(Question)** Complete the following code to implement one iteration of the training process:
- Apply the forward pass on training data and compute the loss
- Apply backpropagation to compute the gradient of the loss with respect to the network parameters
- Apply gradient descent to update the network parameters

<!-- grade_id: cell-064008dd04d210c7, points: 0, grade: false, solution: true -->
```python
class NeuralNetwork(NeuralNetwork):
    def train_iteration(self, X, d, lr=1e-2):
        # Apply forward pass and compute the loss
        y = self.forward(X)
        loss = squared_error(y, d)
        # Apply backpropagation to compute the gradients
        d_W1, d_b1, d_W2, d_b2 = self.backward(X, y, d)
           
        # Apply gradient descent to update the weights
        # YOUR CODE HERE
        self.W1 = self.W1 - lr*d_W1
        self.b1 = self.b1 - lr*d_b1
        self.W2 = self.W2 - lr*d_W2
        self.b2 = self.b2 - lr*d_b2 
        loss = squared_error(y, d)
        return loss
```

<!-- grade_id: train-iteration, points: 0, grade: true, solution: false -->
```python
nn = NeuralNetwork(64, 32, 3)
X = np.random.randn(64, 100)
d = np.random.randint(0, 2, size=(3, 100))
loss = nn.train_iteration(X, d, lr=100)
assert loss >= 0.
```

### 7. Mini-batch training loop

Now, we will implement the main training loop of our neural network.

We will use stochastic gradient descent with mini-batch: the weights will be updated by performing gradient descent on shuffled subsets of training data.

We will train the network for a number of epochs (an epoch is performed when the whole training set has been used with this mini-batch procedure).

**(Question)** Complete the code below to implement the training loop with minibatch stochastic gradient descent.

<!-- grade_id: cell-d38adca6962df05a, points: 0, grade: false, solution: true -->
```python
class NeuralNetwork(NeuralNetwork):
    def fit(self, X, d, batch_size, n_epochs=10, lr=1e-2):
        n_samples = X.shape[1]
        n_batches = (n_samples // batch_size) + 1
        epoch_loss = 0.
        for e in range(n_epochs):
            # Shuffle dataset
            permutation = np.random.permutation(n_samples)
            X, d = X[:, permutation], d[:, permutation]
            # Loop over each batch
            for b in range(0, n_samples, batch_size): # range(start, stop, step)
                # Grab the current batch in `X_batch` and `d_batch`
                X_batch = X[b,b+batch_size]
                d_batch = d[b,b+batch_size]
                # Apply training iteration and update epoch loss
                epoch_loss += self.train_iteration(X_batch, d_batch, lr)
            # Compute average epoch loss and print it
            print(f"Epoch loss: {epoch_loss/n_epochs}")
            
```

### 8. Train the network on the MNIST dataset

The MNIST dataset is composed of 70000 greyscale images of handwritten digits: 60000 images for training and 10000 for testing.

It is included in the `mnist.tgz` archive provided with this TP.

```python
!tar xvzf ./mnist.tgz
images_train = np.load('./mnist-train-images.npy')
labels_train = np.load('./mnist-train-labels.npy')
images_test = np.load('./mnist-test-images.npy')
labels_test = np.load('./mnist-test-labels.npy')
```

Output:
```
mnist-test-images.npy
mnist-test-labels.npy
mnist-train-images.npy
mnist-train-labels.npy
```

**(Question)** Reshape the images into vectors and normalize the pixel values between 0 and 1. Convert the labels into one-hot vectors (*i.e.* vectors full of 0 and with only a 1 for the corresponding class). Store the results into `X_train`, `y_train`, `X_test` and `y_test` variables. Make sure to reshape to the following:
- Input data: `n_features x n_samples`
- Labels: `n_classes x n_samples`

<!-- grade_id: cell-458517fdffea414b, points: 0, grade: false, solution: true -->
```python
X_train = images_train.reshape((images_train.shape[0], (images_train.shape[1]*images_train.shape[2])))
X_test = images_test.reshape((images_test.shape[0], (images_test.shape[1]*images_test.shape[2])))
X_train = X_train.astype('float32')
X_test = X_test.astype('float32')
X_train /= 255
X_test /= 255
n_classes = 10
def arr_one(index):
    arr = np.zeros(n_classes)
    arr[index] = 1
    return arr
y_train = np.array([arr_one(ind) for ind in labels_train]).T
y_test = np.array([arr_one(ind) for ind in labels_test]).T
```

<!-- grade_id: mnist-dataset, points: 0, grade: true, solution: false -->
```python
assert np.all(X_train >= 0.) and np.all(X_train <= 1.)
assert np.all(X_test >= 0.) and np.all(X_test <= 1.)
assert np.all(np.unique(y_train) == np.array([0., 1.])) 
assert np.all(np.unique(y_test) == np.array([0., 1.]))
assert np.all(np.sum(y_train, axis=0) == 1.)
assert np.all(np.sum(y_test, axis=0) == 1.)
```

**(Question)** Initialize a neural network for MNIST with 32 hidden units and train it for 10 epochs with a batch size of 512.

<!-- grade_id: cell-3cc325c6ca9e4632, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

**(Question)** Compute the classification accuracy on the train and test sets. To do that, you can use the predict function and compare them with the original labels (*i.e.* without one-hot encoding).

<!-- grade_id: cell-7b50de5211771a12, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: accuracy, points: 0, grade: true, solution: false -->
```python
assert(0. <= accuracy_train <= 1.)
assert(0. <= accuracy_test <= 1.)
```

**(Question)** Compute and plot the confusion matrix for the test set. Which are the most difficult classes? Show some examples of misclassified images.

<!-- grade_id: cell-a412a782fabb39c4, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

**(Question)** Play around with hyperparameters of the model. What happens when the batch size if very small? And very large?

<!-- grade_id: cell-b6db4b1e04482447, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-dd69c2dd2d0a35d2, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

### 9. (BONUS) Extension to softmax and categorical cross-entropy

**(Question)** Extend your neural network model to use a softmax activation function for the output layer, and a categorical cross-entropy loss.
You can also experiment with the reLU activation for the hidden layer.

*Hint:* recall the partial derivatives formulation from logistic regression, and optimize the backpropagation for the output layer accordingly.

```python

```

**(Question)** Extend your neural network model to handle more than one hidden layer.

```python

```